Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

## CSL2050: Pattern Recognition and Machine Learning<br>
Programming Assignment-2<br>
Spring 2025<br>

---
## KNN Classification
In this assignment, we will implement K-nearest neighbour for Handwritten Digit Classification problem. For handwritten digit classification problem, we shall use a small subset of MNIST Dataset. 

KNN algorithm has three important modules: (i) Distance Computation and (ii) Finding K nearest neigbour to a point, (iii) Predicting the label for a test point. Let us first implement them. 

**Problem 1.01:** Write a function to compute distance between two points. The function should compute either Eucleadean or Manhattan distance based on the argument passed in the function. (2 points) 

<!-- grade_id: calDist, points: 0, grade: false, solution: true -->
```python
import math
def calDist(x,y,metric):
    """Compute the distance (either euclidean or manhattan)
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print a value exactly or close to (with absolute tolerance = 0.01) 1.41 for calDist([1,1],[0,0],'euclidean'). 
Your function should print a value exactly or close to (with absolute tolerance = 0.01) 3 for calDist([1,1],[0,0],'manhatten'). 
It should support higher dimension inputs. 

```python
calDist([1,1],[0,0],'euclidean')
```

```python
calDist([1,1],[0,0],'manhattan')
```

```python
calDist([1,1,2],[0,0,1],'manhattan')
```

<!-- grade_id: CalDistTest, points: 2, grade: true, solution: false -->
```python
"""Check that calDist returns the correct output for several inputs"""
import math
assert math.isclose(calDist([1,1],[0,0],'euclidean'), 1.414, abs_tol=0.01)
assert math.isclose(calDist([1,1],[0,0],'manhattan'), 2, abs_tol=0.01)
assert math.isclose(calDist([1,1,2],[0,0,1],'manhattan'), 3, abs_tol=0.01)
```

**Problem-1.02:** Write a function that returns the k nearest neighbors to a given point from a set of points, along with their corresponding indices in the set.

<!-- grade_id: cell-7b77aeca027bdc3e, points: 0, grade: false, solution: true -->
```python
import numpy as np

def k_nearest_neighbors(p, points, k, metric):
    """
    Finds the k nearest neighbors of point `p` in the dataset `points` using the specified distance metric.
    
    Parameters:
    - p (list): The point (e.g., [x1, y1, z1, ...]).
    - points (list of lists): The set of points (e.g., [[x1, y1, z1], [x2, y2, z2], ...]).
    - k (int): The number of nearest neighbors to find.
    - metric (str): The distance metric, either 'euclidean' or 'manhattan'.
    
    Returns:
    - list: A list of the k nearest neighbors (each neighbor is a list representing a point) 
            and their index in `points`.
    """
     # YOUR CODE HERE
     raise NotImplementedError()

```

```python
points = [[1, 2], [3, 4], [5, 6], [7, 8], [2, 3], [6, 7]]
query_point = [1, 2]
k = 3
a,b=k_nearest_neighbors(query_point, points, k, 'euclidean')
print(a)
print(b)
```

```python
points = [[1, 2], [3, 4], [5, 6], [7, 8], [2, 3], [6, 7]]
query_point = [6, 7]
k = 3
a,b=k_nearest_neighbors(query_point, points, k, 'manhattan')
print(a)
print(b)
```

<!-- grade_id: k_nearest_neighborsTest, points: 2, grade: true, solution: false -->
```python
"""Check that k_nearest_neighbors returns the correct output for several inputs using hidden test cases"""
```

**Problem-1.03**: Write a function that predicts the labels for validation data using the K-Nearest Neighbors (KNN) algorithm. The function should take the following arguments:

    X_train: A numpy array of feature data for training.
    y_train: A numpy array of labels for the training data.
    testPoint: A single validation or test data sample to predict the label for.
    k: The number of nearest neighbors to consider for the prediction.
    distance_metric: The type of distance metric to use ('euclidean' or 'manhattan').

The function should return the predicted labels for the validation data sample based on the KNN algorithm.

<!-- grade_id: knn_predict, points: 0, grade: false, solution: true -->
```python
import numpy as np

def knn_predict(X_train, y_train, testPoint, k, distance_metric='euclidean'):
  """
  Predicts labels for validation data using the K-Nearest Neighbors algorithm.

  Args:
    X_train: Training data features (numpy array).
    y_train: Training data labels (numpy array).
    testPoint: Validation or test data sample (numpy array).
    k: Number of nearest neighbors to consider.
    distance_metric: Distance metric to use ('euclidean' or 'manhattan').

  Returns:
    predictions: List of predicted labels for the validation data.
  """
  # YOUR CODE HERE
  raise NotImplementedError()

# Example usage:
# Assuming you have functions for k_nearest_neighbors and necessary imports
#X_train = np.array([[1, 2], [2, 3], [3, 4], [6, 7], [7, 8]])
#y_train = np.array([0, 0, 1, 1, 1])  # Labels for the training data
#testPoint = np.array([2, 2])  # Test point for prediction
#k = 3  # Number of neighbors
#predicted_label = knn_predict(X_train, y_train, testPoint, k, 'euclidean')
```

<!-- grade_id: knn_predictTest, points: 4, grade: true, solution: false -->
```python
"""Check that knn_predict returns the correct output for several inputs using hidden test cases"""

```

Now, we shall use our learning for classifying handwritten digits into one of the classes (0 to 9). Complete the following code so that accuracy on validation data can be computed. You should use knn_predict function you have already written to complete this.

**Problem-1.04**: Review the provided code and complete it by invoking the knn_predict function for each sample in the validation dataset. The predictions list should store the predicted labels for each sample in the validation set.

<!-- grade_id: knn, points: 5, grade: true, solution: true -->
```python
!pip install tqdm
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
from tqdm import tqdm

# Step 1: Load the mnist_student.csv dataset
# Load the data from the provided mnist_student.csv file
data = pd.read_csv('mnist_student.csv')

# Separate features (X) and labels (y)
X = data.drop('label', axis=1)  # Features (28x28 pixel images flattened)
y = data['label']  # Labels

# Normalize the pixel values to be between 0 and 1
X = X / 255.0

# Step 2: Split the dataset into training and validation sets
X_train, X_val, y_train, y_val = train_test_split(X, y, train_size=0.8, random_state=42)
X_train=X_train.values
y_train=y_train.values
X_val=X_val.values
y_val=y_val.values
k=5
predictions=[]

## Complete this code by calling knn_predict for each sample of validation data and
## predictions should store predicted labels for each of the validation set samples
# YOUR CODE HERE
raise NotImplementedError()

# Compute accuracy on validation set
accuracy = accuracy_score(y_val, predictions)
print(f'Validation Accuracy: {accuracy * 100:.2f}%')

# You can also display some images and their predicted labels here

```

Next, we provide you a helper code that you can use to visualize the result of your implementation. It will randomly choose 5 samples from validation set and show true (actual) and predicated labels. 

```python
import matplotlib.pyplot as plt
import numpy as np

# Assuming `predictions` is the list of predicted labels for the validation set

# Step 3: Display some of the predicted samples
def plot_samples(X_val, y_val, predictions, num_samples=5):
    # Randomly select `num_samples` to display
    random_indices = np.random.choice(len(X_val), size=num_samples, replace=False)
    
    fig, axes = plt.subplots(1, num_samples, figsize=(15, 5))
    for i, idx in enumerate(random_indices):
        # Get the sample image (reshape from 1D array to 28x28)
        img = X_val[idx].reshape(28, 28)
        true_label = y_val[idx]
        predicted_label = predictions[idx]
        
        # Plot the image
        axes[i].imshow(img, cmap='gray')
        axes[i].set_title(f"True: {true_label}\nPred: {predicted_label}")
        axes[i].axis('off')
    
    plt.show()

# Call the function to display 5 random samples
plot_samples(X_val, y_val, predictions, num_samples=5)

```

Review, understand and try out this code which draws confusion matrix on validation set. 

```python
!pip install seaborn
import numpy as np
import matplotlib.pyplot as plt
from sklearn.metrics import confusion_matrix
import seaborn as sns

# Assuming `predictions` contains the predicted labels and `y_val` contains the true labels

# Step 1: Compute the confusion matrix
cm = confusion_matrix(y_val, predictions)

# Step 2: Plot the confusion matrix
def plot_confusion_matrix(cm, class_names):
    plt.figure(figsize=(10, 7))
    sns.heatmap(cm, annot=True, fmt='d', cmap='Blues', xticklabels=class_names, yticklabels=class_names)
    plt.title('Confusion Matrix')
    plt.xlabel('Predicted Labels')
    plt.ylabel('True Labels')
    plt.show()

# Step 3: Get class names (if not already available)
class_names = [str(i) for i in range(10)]  # For MNIST dataset (digits 0-9)

# Step 4: Call the function to plot the confusion matrix
plot_confusion_matrix(cm, class_names)

```

**Problem-1.05:** What value of k gave you best validation accuracy. How many values you tried? 

**Problem-1.06:** Write your observation about failure cases from visual inspection of results.

**Problem-1.07:** What did you observe from the confusion matrix? 

**Problem-1.08:** For a fix k, which distance metric gave you better result? 

**Problem-1.09:** How did you handle ties in the k-NN algorithm? For instance, if the three nearest neighbors belong to three different classes, how is the final prediction determined? Which part of your code addresses this issue? Please point it out.

**Problem-1.10: Ethical Reflection and Acknowledgments (Mandatory Question)**
(i) List all collaborators, references, or resources you used. If none, write "NA."

(ii) Estimate the percentage of the code you wrote yourself.

(iii) Reflect on your ethical practices (Yes/No):

    (a) Did you avoid copying code without understanding it?
    (b) Did you properly cite all resources and collaborators?
        