Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

## CSL2050: Pattern Recognition and Machine Learning<br>
Programming Assignment-2<br>
Spring 2025<br>

## Linear Regression
In this assignment, we will explore linear regression by implementing fundamental modules such as Mean Squared Error (MSE), Gradient Descent, and Prediction. These building blocks will then be utilized to experiment with a real-world dataset for solving a linear regression problem. Please ensure that you solve the problems sequentially for a structured learning experience. 

**Problem-2.01:** Write a function mean_squared_error(y_true, y_pred) that calculates the Mean Squared Error (MSE) between the true values y_true and predicted values y_pred.

<!-- grade_id: mean_squared_error, points: 0, grade: false, solution: true -->
```python
def mean_squared_error(y_true, y_pred):
    """
    Calculate the Mean Squared Error between true and predicted values.

    Args:
        y_true (numpy.ndarray): True values of shape (n_samples,).
        y_pred (numpy.ndarray): Predicted values of shape (n_samples,).

    Returns:
        float: Mean Squared Error.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: mean_squared_error_test, points: 2, grade: true, solution: false -->
```python

# Visible Test Case
import numpy as np
y_true = np.array([1.0, 2.0, 3.0])
y_pred = np.array([1.0, 2.0, 3.0])
assert mean_squared_error(y_true, y_pred) == 0, "Test Case 1 Failed"

```

**Problem-2.02:** Write a function ols_coefficients(X, y) that calculates the coefficients for a simple linear regression problem using the Ordinary Least Squares (OLS) formula. Note: Coefficient in hyperplane \theta_0+\theta_1x_1+\theta_2x_2 are \theta_0, \theta_1 and \theta_2. 

<!-- grade_id: ols_coefficients, points: 0, grade: false, solution: true -->
```python
def ols_coefficients(X, y):
    """
    Calculate the OLS coefficients for linear regression.

    Args:
        X (numpy.ndarray): Feature matrix of shape (n_samples, n_features).
        y (numpy.ndarray): Target vector of shape (n_samples,).

    Returns:
        numpy.ndarray: Coefficients of shape (n_features + 1,).
                       The first value is the intercept, and the rest are feature coefficients.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: ols_coefficients_test, points: 2, grade: true, solution: false -->
```python
!pip install pytest
import pytest
# Test case 1: Single feature, no intercept
def test_case_1():
    X = np.array([[1], [2], [3]])
    y = np.array([2, 4, 6])
    expected = np.array([0., 2.])  # y = 2x
    np.testing.assert_almost_equal(ols_coefficients(X, y), expected, decimal=6)

# Test case 2: Single feature, with intercept
def test_case_2():
    X = np.array([[1], [2], [3]])
    y = np.array([3, 5, 7])
    expected = np.array([1., 2.])  # y = 1 + 2x
    np.testing.assert_almost_equal(ols_coefficients(X, y), expected, decimal=6)

test_case_1()
test_case_2()
```

**Problem-2.03:** Implement a function gradient_descent(X, y, lr, epochs) to perform gradient descent for linear regression.

Given a feature matrix X of shape (n_samples,n_features) a target vector y of shape (n_samples,), a learning rate lr, and the number of iterations epochsepochs, your task is to iteratively update the coefficients θ to minimize the mean squared error between the predicted and actual target values.

The function should return the optimized coefficients θ, which are of shape (n_features+1,). (Note the line/hyperplan has equations like θ_0 + θ_1x_1 + θ_2x_2 + ....+θ_nx_n=0.)

<!-- grade_id: gradient_descent, points: 0, grade: false, solution: true -->
```python
def gradient_descent(X, y, lr, epochs):
    """
    Perform gradient descent for linear regression.

    Args:
        X (numpy.ndarray): Feature matrix of shape (n_samples, n_features).
        y (numpy.ndarray): Target vector of shape (n_samples,).
        lr (float): Learning rate.
        epochs (int): Number of iterations.

    Returns:
        numpy.ndarray: Optimized coefficients of shape (n_features+1).
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: gradient_descent_test, points: 2, grade: true, solution: false -->
```python
# Test case 1: Simple case with one feature and bias term
def test_case_1():
    X = np.array([[1], [2], [3]])
    y = np.array([2, 4, 6])  # Perfectly linear relationship y = 2x
    lr = 0.01
    epochs = 1000
    theta = gradient_descent(X, y, lr, epochs)
    # Assert bias and slope
    assert np.abs(theta[0] - 0) < 0.5, f"term {theta[0]} is not close enough to 0"
    assert np.abs(theta[1] - 2) < 0.5, f"term {theta[1]} is not close enough to 2"
test_case_1()
```

**Problem-2.04:** Write a function predict(X, theta) that predicts the target values y for each sample of a given feature matrix X and the learned coefficients θ.

<!-- grade_id: predict, points: 0, grade: false, solution: true -->
```python
def predict(X, theta):
    """
    Predict target values for given features and coefficients.

    Args:
        X (numpy.ndarray): Feature matrix of shape (n_samples, n_features).
        theta (numpy.ndarray): Coefficients of shape (n_features,).

    Returns:
        numpy.ndarray: Predicted values of shape (n_samples,).
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: predict_test, points: 2, grade: true, solution: false -->
```python
import numpy as np
# Test case 1: Single feature, single sample
def test_case_1():
    X = np.array([[2]])  # Single sample with one feature
    theta = np.array([1, 2])  # Bias = 1, Coefficient = 2
    expected = np.array([5])  # Prediction: 1 + 2*2 = 5
    np.testing.assert_almost_equal(predict(X, theta), expected, decimal=6)

# Test case 2: Single feature, multiple samples
def test_case_2():
    X = np.array([[1], [2], [3]])  # Three samples, one feature
    theta = np.array([0.5, 1.5])  # Bias = 0.5, Coefficient = 1.5
    expected = np.array([2.0, 3.5, 5.0])  # Predictions: 0.5 + 1.5*X
    np.testing.assert_almost_equal(predict(X, theta), expected, decimal=6)

test_case_1()
test_case_2()
```

**Problem-2.05:** Putting it all together, now we can write a function that leverages the previously implemented functions, including Mean Squared Error (MSE), gradient descent, and prediction, and perform linear regression using gradient descent (assuming MSE loss). 

<!-- grade_id: LinearRegressionUsingGD, points: 0, grade: false, solution: true -->
```python
def LinearRegressionUsingGD(X,y,lr, epochs):
  theta=gradient_descent(X, y, lr, epochs)
  predictions=predict(X,theta)
  #Write appropriate return statement
  # YOUR CODE HERE
  raise NotImplementedError()
```

```python
# Example Usage
import numpy as np

X = np.array([[1], [2], [3]])  
y = np.array([1, 2, 3])

theta, predictions = LinearRegressionUsingGD(X,y,lr=0.01, epochs=1000)

print("Coefficients:", theta)
print("Predictions:", predictions)

```

<!-- grade_id: LinearRegressionUsingGD_test, points: 6, grade: true, solution: false -->
```python
import numpy as np

# Test Case 1: Simple linear regression with a single feature
def test_case_1():
    X = np.array([[1], [2], [3]])  # Feature matrix
    y = np.array([1, 2, 3])  # Target values

    theta, predictions = LinearRegressionUsingGD(X,y,lr=0.01, epochs=1000)
    
    # Calculate residuals
    residuals = np.abs(predictions - y)

    # Assert that residuals are very small (e.g., less than 0.01)
    np.testing.assert_array_less(residuals, np.full_like(residuals, 0.4))
    
test_case_1()

```

**Problem-2.06:** Leverage your implementation and choose any real word regression dataset of your choice. Split the data into 80-10-10% of train-val-test. Report train and test MSE of your linear regression model. 

**Problem-2.07:** Which dataset you have chosen for Problem-2.6. Expalin briefly about the task in the dataset. How does your test error changes with training datasize, explore and report quantitative analysis. 

**Problem-2.08:** Ethical Reflection and Acknowledgments (Mandatory Question) (i) List all collaborators, references, or resources you used. If none, write "NA."

(ii) Estimate the percentage of the code you wrote yourself.

(iii) Reflect on your ethical practices (Yes/No):

(a) Did you avoid copying code without understanding it?
(b) Did you properly cite all resources and collaborators?
    